<?xml version="1.0" encoding="UTF-8"?>
<quiz>
<question type="stack">
    <name>
        <text>All-or-Nothing graded set equality question</text>
    </name>
    <questiontext format="html">
        <text><![CDATA[<p>Given distance A &lt;-&gt; B = 2 and distance B &lt;-&gt; C = 4, return the k-Means cluster with k = 1.</p>
<p>[[input:received]] [[validation:received]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
        <text></text>
    </generalfeedback>
    <idnumber></idnumber>
    <defaultgrade>1.0</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <stackversion>
      <text>2025040100</text>
    </stackversion>
    <questionvariables>
        <text>expected: { A,B,C }</text>
    </questionvariables>
    <specificfeedback format="html">
        <text><![CDATA[<p>[[feedback:prt1]]</p>]]></text>
    </specificfeedback>
    <questionnote format="html">
        <text><![CDATA[<p>{@expected@}</p>]]></text>
    </questionnote>
    <questiondescription format="html">
        <text><![CDATA[]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
        <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
        <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
        <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
        <name>received</name>
        <type>algebraic</type>
        <tans>expected</tans>
        <boxsize>37</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>1</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>2</showvalidation>
        <options></options>
    </input>

    <prt>
        <name>prt1</name>
        <value>1.0</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description></description>
            <answertest>Sets</answertest>
            <sans>received</sans>
            <tans>expected</tans>
            <testoptions></testoptions>
            <quiet>1</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1.0</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>prt1-expected-correct</trueanswernote>
            <truefeedback format="html">
              <text><![CDATA[]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0.0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>prt1-expected-wrong</falseanswernote>
            <falsefeedback format="html">
              <text><![CDATA[]]></text>
            </falsefeedback>
        </node>
    </prt>
</question>
<question type="stack">
    <name>
        <text>All-or-Nothing graded set equality question as subquestion</text>
    </name>
    <questiontext format="html">
        <text><![CDATA[<p>Given distance A &lt;-&gt; B = 2 and distance B &lt;-&gt; C = 4, return the k-Means cluster with k = 1.</p>
<p>[[input:ese_received]] [[validation:ese_received]]</p>]]></text>
    </questiontext>
    <generalfeedback format="html">
        <text></text>
    </generalfeedback>
    <idnumber></idnumber>
    <defaultgrade>1.0</defaultgrade>
    <penalty>0</penalty>
    <hidden>0</hidden>
    <stackversion>
      <text>2025040100</text>
    </stackversion>
    <questionvariables>
        <text>ese_expected: { A,B,C }</text>
    </questionvariables>
    <specificfeedback format="html">
        <text><![CDATA[
<p>[[feedback:ese_prt]]</p>]]></text>
    </specificfeedback>
    <questionnote format="html">
        <text><![CDATA[]]></text>
    </questionnote>
    <questiondescription format="html">
        <text><![CDATA[]]></text>
    </questiondescription>
    <questionsimplify>1</questionsimplify>
    <assumepositive>0</assumepositive>
    <assumereal>0</assumereal>
    <prtcorrect format="html">
        <text></text>
    </prtcorrect>
    <prtpartiallycorrect format="html">
        <text></text>
    </prtpartiallycorrect>
    <prtincorrect format="html">
        <text></text>
    </prtincorrect>
    <decimals>.</decimals>
    <scientificnotation>*10</scientificnotation>
    <multiplicationsign>dot</multiplicationsign>
    <sqrtsign>1</sqrtsign>
    <complexno>i</complexno>
    <inversetrig>cos-1</inversetrig>
    <logicsymbol>lang</logicsymbol>
    <matrixparens>[</matrixparens>
    <variantsselectionseed></variantsselectionseed>
    <input>
        <name>ese_received</name>
        <type>algebraic</type>
        <tans>ese_expected</tans>
        <boxsize>37</boxsize>
        <strictsyntax>1</strictsyntax>
        <insertstars>0</insertstars>
        <syntaxhint></syntaxhint>
        <syntaxattribute>0</syntaxattribute>
        <forbidwords></forbidwords>
        <allowwords></allowwords>
        <forbidfloat>1</forbidfloat>
        <requirelowestterms>0</requirelowestterms>
        <checkanswertype>0</checkanswertype>
        <mustverify>1</mustverify>
        <showvalidation>2</showvalidation>
        <options></options>
    </input>

    <prt>
        <name>ese_prt</name>
        <value>1.0</value>
        <autosimplify>1</autosimplify>
        <feedbackstyle>1</feedbackstyle>
        <feedbackvariables>
            <text></text>
        </feedbackvariables>
        <node>
            <name>0</name>
            <description></description>
            <answertest>Sets</answertest>
            <sans>ese_received</sans>
            <tans>ese_expected</tans>
            <testoptions></testoptions>
            <quiet>1</quiet>
            <truescoremode>=</truescoremode>
            <truescore>1.0</truescore>
            <truepenalty></truepenalty>
            <truenextnode>-1</truenextnode>
            <trueanswernote>ese_prt-ese_expected-correct</trueanswernote>
            <truefeedback format="html">
              <text><![CDATA[]]></text>
            </truefeedback>
            <falsescoremode>=</falsescoremode>
            <falsescore>0.0</falsescore>
            <falsepenalty></falsepenalty>
            <falsenextnode>-1</falsenextnode>
            <falseanswernote>ese_prt-ese_expected-wrong</falseanswernote>
            <falsefeedback format="html">
              <text><![CDATA[]]></text>
            </falsefeedback>
        </node>
    </prt>
</question>
</quiz>
